---
title: '2'
description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise
key: bbe2a6e243
xp: 100
```

nicka

`@instructions`
dsadas

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: TabExercise
key: beae7c65c7
xp: 100
```

nick

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

***

```yaml
type: NormalExercise
key: aea726526d
xp: 50
```

`@instructions`
nick

`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

***

```yaml
type: NormalExercise
key: 19d1bb6290
xp: 50
```

`@instructions`
nick

`@hint`


`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
