---
  title: "2"
  description: ""
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: bbe2a6e243   
```


nicka


`@instructions`
dsadas

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: beae7c65c7   
```


nick


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


***



```yaml
type: NormalExercise 
xp: 50 
key: aea726526d   
```





`@instructions`
nick

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 50 
key: 19d1bb6290   
```





`@instructions`
nick

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


