---
  title: "Test"
  description: "Test"
---

## 2

```yaml
type: NormalExercise 
lang: python
xp: 100 
skills: 2
key: b414e5d252   
```


This is the assignment text. It should help provide students with the background information needed.
The instructions that follow should be in bullet point form dawith clear guidance for what is expected.


`@instructions`
- Instruction 1
- Instruction 2
- Instruction 3

`@hint`
- Here is the hint for this setup problem. 
- It should get students 50% of the way to the correct answer.
- So don't provide the answer, but don't just reiterate the instructions.
- Typically one hint per instruction is a sensible amount.

`@pre_exercise_code`

```{python}
# Load datasets and packages here.
```

`@sample_code`

```{python}
# Your
# sample
# code
# should
# be
# ideally
# 10 lines or less,
# with a max
# of 16 lines.
```

`@solution`

```{python}
# Answer goes here
# Make sure to match the comments with your sample code
# to help students see the differences from solution
# to given.
```

`@sct`

```{python}
# Update this to something more informative.
success_msg("Some praise! Then reinforce a learning objective from the exercise.")
```

---

## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 76c9e3f677   
```

`@projector_key`
d453f4baed6ff440bec6d1c0bc9a9df4
---

## 1

```yaml
type: NormalExercise 
xp: 100 
key: 960e1e1889   
```


Hey Nick ok this  live
da
- instruction
- instruction
- instruction- instruction
- instruction
- instruction
- instruction
- instruction
- instruction
- instruction

nick

dsa
das
das
d
asd
as
da


`@sct`

```{python}
Hey Nick ok this  live

- instruction
- instruction
- instruction- instruction
- instruction
- instruction
- instruction
- instruction
- instruction
- instruction

nick

dsa
das
das
d
asd
as
da
```

---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: a6c5670169   
```





---

## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: d1ed79bc76   
```





`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}

```


***



```yaml
type: NormalExercise 
xp: 35 
key: 932efbe8b5   
```





`@instructions`
2

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 35 
key: 7623fd22c7   
```





`@instructions`
1

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```



***



```yaml
type: NormalExercise 
xp: 30 
key: 710d0f4bba   
```





`@instructions`
3

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


